```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment2.ipynb")
```

# Otter-Grader Tutorial

This notebook is part of the Otter-Grader tutorial. For more information about Otter, see our [documentation](https://otter-grader.rtfd.io).

```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 2:** Write an infinite generator of the Fibonacci sequence `fiberator` that is *not* recursive.

```python
def fiberator():
    # BEGIN SOLUTION
    yield 0
    yield 1
    x, y = 0, 1
    while True:
        x, y = y, x + y
        yield y
    # END SOLUTION
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

**Question 5:** What do you notice about the bar plot?

_Type your answer here, replacing this text._

**SOLUTION:** mint is the highest...?

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2 (points: None)
```
>>> f = fiberator()
>>> assert next(f) == 0
>>> assert next(f) == 1
```
```
>>> f = fiberator()
>>> assert next(f) == 0
>>> assert next(f) == 1
>>> assert next(f) == 1
>>> assert next(f) == 2
>>> assert next(f) == 3
>>> assert next(f) == 5
>>> assert next(f) == 8
>>> assert next(f) == 13
>>> assert next(f) == 21
```
